# Homework 3


## Problem 1 (Reocurring Characters: dictionaries) 

For a given string s, return the first reoccuring character in the string.  Try to get an approach that is $O(n)$, where $n$ is the number of characters in the string. If there is no such character, return None. 
    
    -- Input:"ABCA"
    -- Output:"A"

<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
def problem1(s):
    
    # YOUR CODE HERE
    s_list = [i for i in s]
    s_dict = {i: 0 for i in s_list}
    
    for i in s_list:
        s_dict[i] += 1
        if s_dict[i] == 2:
            return(i)
```

<!-- grade_id: cell-2853176972493b6e, points: 10, grade: true, solution: false -->
```python
assert problem1("ABCA") == "A"
```

<!-- grade_id: cell-b16dd3056b5a9f95, points: 10, grade: true, solution: false -->
```python
assert problem1("ABC") == None
```

<!-- grade_id: cell-ab1b1d0a5f7b1fce, points: 10, grade: true, solution: false -->
```python
assert problem1("ABCBA") == "B"
```

## Problem 2 (String Methods):

We have learnt that strings are immutable objects. In the following, you are asked to write a code to "mutate" a string. In particular, you are given a string, a position number and a character. You should return a different string such that this string is the same as the input string except the character in the position is changed to the given character. 

Example:

    input: 
        original_string = "Hello world!"
        position=6
        character="W"
    ouput: "Hello World!"

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
def problem2(original_string, position, character):
    
    # YOUR CODE HERE
    return_string = original_string[:position] + character + original_string[position+1:]
    return(return_string)
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 10, grade: true, solution: false -->
```python
# Test cases for Problem Set 4
assert problem2("Hello world!", 6, "W") == "Hello World!"
```

<!-- grade_id: cell-402bf89f76c4bd42, points: 10, grade: true, solution: false -->
```python
assert problem2("Hello world!", 11, ".") == "Hello world."
```

```python

```